<center>
<img src="logo.png" height="900"> 
</center>


#  Распределения

В этом задании мы немного поработаем в python с разными случайными величинами. Делать это задание необязательно, но рекомендуется. 

```python
import numpy as np
import pandas as pd

import scipy.stats as sts
import matplotlib.pyplot as plt
import seaborn as sns

plt.style.use('ggplot')  # стиль для графиков
%matplotlib inline
```

## Упражнение 1 

> Нормальность восстановлена, мы на орбите Магратеи (Автостопом по галактике) 

__а)__ Задайте генератор для случайной величины имеющей нормальное распределение: $X \sim N(4, 10^2)$. 

<!-- grade_id: cell-21e561bdc8c563e2, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# will the code be with you

norm_rv = ... 
```

__б)__ Найдите вероятности $P(X > 4)$, $P(X \in [2; 20])$, $P(X < −5)$. Запишите их в переменные `p1, p2, p3`. 

<!-- grade_id: cell-14d3057847804744, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# will the code be with you

p1 = ... 
p2 = ...
p3 = ...

```

<!-- grade_id: cell-ebe87628dc40aa0d, points: 1, grade: true, solution: false -->
```python
# Проверка корректно ли вы нашли вероятности :) 
# Задание необязательное, поэтому все тесты открытые

assert p1 == 0.5
assert np.abs(p2 - 0.5244604) < 1e-5
assert np.abs(p3 - 0.18406012) < 1e-5
```

__в)__ Найдите число $a$ такое, что $P(X > a) = 0.3$.

<!-- grade_id: cell-42289d9eda89b327, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# will the code be with you

a = ...

```

<!-- grade_id: cell-72a1fc8db7e3a18e, points: 1, grade: true, solution: false -->
```python
assert np.abs(a - 9.2440051) < 1e-5
```

__г)__ Найдите число $b$ такое, что $P(X \in [4 - b; 4 + b]) = 0.5$.

<!-- grade_id: cell-ebe9ef1162005659, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# will the code be with you

b = ...

```

Output:
```
0.5
```

<!-- grade_id: cell-a918707ad64c1dd2, points: 1, grade: true, solution: false -->
```python
assert np.abs(b - 6.7448975) < 1e-5
```

__д)__ Сгенерируйте выборку размера $100$. Постройте по этой выборке гистограмму. На том же рисунке изобразите плотность распределения нормальной случайной величины. 

<!-- grade_id: cell-62774151762c127c, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# will the code be with you


```

__е)__ Оцените эмпирическую функцию распределения. Изобразите её и теоретическую функцию распределения на графике. 

<!-- grade_id: cell-99ec9324e1bf3f73, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# will the code be with you


```

## Упражнение 2 

> Звёзды лучше видны с крыши, полезай и проверь сам. Ты так ждал этот знак свыше, и отметил его как спам. (Дайте танк) 

Пусть количество писем со спамом, которое пришло к нам на почту, имеет распределение Пуассона. Предположим, что вы получаете в среднем три спам-письма в день. Какова доля дней, в которые вы получаете пять или больше спам-писем?

<!-- grade_id: cell-e9335d7a188e9bb6, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# will the code be with you

p = ...


```

<!-- grade_id: cell-98ceec526b48dda9, points: 1, grade: true, solution: false -->
```python
assert np.abs(p - 0.1847367) < 1e-5
```

## Упражнение 3 

Во время ЧЕ по футболу 2008 года и ЧМ 2010 года Осьминог Пауль занимался прогнозированием побед (после он ушёл в финансовую аналитику и IB). Осьминог дал верные прогнозы в 12 случаях из 14. Если предположить, что Пауль выбирает победителя наугад, какова вероятность получить 12 верных прогнозов из 14?

<!-- grade_id: cell-14be6ed99a0d0875, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# will the code be with you

p = ...


```

<!-- grade_id: cell-2d5479b3d44226ea, points: 1, grade: true, solution: false -->
```python
assert np.abs(p - 0.0055541) < 1e-5
```

Если предположить, что Осьминог правильно выбирает победителя с вероятностью $0.9$, какова вероятность получить тот же результат? 

<!-- grade_id: cell-6b68d3e49ada3496, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# will the code be with you

p = ...


```

<!-- grade_id: cell-f594071da5301c73, points: 1, grade: true, solution: false -->
```python
assert np.abs(p - 0.2570108) < 1e-5
```