<center>
<img src="logo.png" height="900"> 
</center>


#  Гипотеза о равенстве средних

В этом задании мы попробуем сравнить между собой мощность двух процедур для проверки гипотезы о равенстве средних. __Это задание не является обязательным. Его можно не решать.__ 

-----------

```python

```

```python
import numpy as np
import pandas as pd
import seaborn as sns
import matplotlib.pyplot as plt
from scipy import stats

plt.style.use('ggplot')
%matplotlib inline
```

Поговорим про мощность различных процедур. Мы с вами обсудили, что гипотезу о равенстве средних можно проверить с помощью следующей процедуры. 

### Процедура 1:

1. Собираем выборки $X_1, \ldots, X_n$ и $Y_1, \ldots, Y_n$;
2. Находим значение статистики

$$
z_{obs} = \frac{\bar x - \bar y}{\sqrt{\frac{s_x^2}{n_x} + \frac{s_y^2}{n_y}}};
$$

3. Говорим, что по ЦПТ $z_{obs} \overset{asy}{\sim} N(0,1);$
4. Находим критическое значение $z_{1 - \frac{\alpha}{2}}$;
5. Если мы видим, что $|z_{obs}| <  z_{1 - \frac{\alpha}{2}}$, мы говорим, что гипотеза не отвергается. 



Ту же саму гипотезу можно попробовать проверить с помощью другого алгоритма, основанного на доверительных интервалах.

### Процедура 2: 

1. Собираем выборки $X_1, \ldots, X_n$ и $Y_1, \ldots, Y_n$;
2. Находим $\bar x$ и $\bar y$;
3. Пользуясь ЦПТ и зная, что $\bar x \overset{asy}{\sim} N \left(\mu_1,\frac{s^2_x}{n_x} \right)$ и $\bar y \overset{asy}{\sim} N\left(\mu_2,\frac{s^2_y}{n_y}\right)$ строим для $\mu_1$ и $\mu_2$ доверительные интервалы;
4. Если доверительные интервалы пересеклись, говорим, что гипотеза не отвергается. 


Вроде бы вторая процедура выглядит довольно естественно, однако ей никто не пользуется. Дело в том, что для одинаковых ошибок первого рода, $\alpha$, ошибка второго рода, $\beta$, для процедуры, основанной на доверительных интервалах, окажется выше. Задание состоит в том, чтобы это увидеть.

> Для простоты будем дальше предполагать, что $\bar x > \bar y$. Также будем считать, что обе дисперсии известны и равны единице, $\sigma^2_x = \sigma^2_y = 1$. Объёмы выборок одинаковы, $n_x = n_y = n$. Альтернатива односторонняя, то есть наблюдаемое значение статистики нужно искать как $z_{1 - \alpha}$. 

__[а]__  Найдите мощность первой процедуры (единица минус ошибка второго рода). Запишите её через функцию распределения стандартной нормальной величины $\Phi(x)$. Вбейте её в заглушку ниже.

<!-- grade_id: cell-dd29bcf28c5bbb42, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# will the code be with you
from scipy.stats import norm

def power_1(n_obs, alpha=0.05, mu1=4, mu2=5, sigma=1):
    

# your code here
```

Output:
```
n=100, power=1.00000
n=10, power=0.93542
1.0 0.9354201700477403
```

<!-- grade_id: cell-9a089f0c7a047ead, points: 1, grade: true, solution: false -->
```python
# Ура ура ура все тесты открыты :) 
assert np.abs(power_1(100) - 0.99999) < 1e-4
assert np.abs(power_1(10) - 0.72281) < 1e-4
```

Output:
```
n=100, power=1.00000
n=10, power=0.93542
AssertionError                            Traceback (most recent call last)
Cell In[18], line 3
      1 # Ура ура ура все тесты открыты :) 
      2 assert np.abs(power_1(100) - 0.99999) < 1e-4
----> 3 assert np.abs(power_1(10) - 0.72281) < 1e-4

AssertionError:
```

__[б]__ Выпишите доверительный интервал для $\bar x$ и для $\bar y$.  Правая граница интервала для $\bar x$ должна оказаться правее левой границы для $\bar y$. В таком случае гипотеза не отвергается. 

- Выпишите это неравенство. 
- Выразите из него значение $z_{1 - \frac{\alpha}{2}}$.
- Сравните получившуюся статистку со статистикой из первого пункта. 
- Найдите для неё мощность, запишите её через функцию распределения стандартной нормальной величины, $\Phi(x)$.

<!-- grade_id: cell-95e1c68560ce72f6, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# will the code be with you

def power_2(n_obs, alpha=0.05, mu1=4, mu2=5, sigma=1):
    


# your code here

```

<!-- grade_id: cell-087864869f8bdd18, points: 1, grade: true, solution: false -->
```python
# Ура ура ура все тесты открыты :) 
assert np.abs(power_2(100) - 0.9999) < 1e-4
assert np.abs(power_2(10) - 0.46410) < 1e-4
```

__[в]__  Постройте для получившихся функций графики, сравните их между собой. 


<!-- grade_id: cell-b4b13617e751a7bb, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# will the code be with you

# your code here

```